---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: ddc011c0e5
xp: 100
```

Assign the number 2 to variable a. 

`@instructions`
Assign the number 2 to variable a.

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
a <- 
```

`@solution`
```{r}
a <- 2
```

`@sct`
```{r}

```
